# AM1054 Continuous Assessment 1

```python
# execute this cell before you start the assignment
%matplotlib inline
import numpy as np
import matplotlib.pyplot as plt
```

Due Thursday 25 October 2018, 4pm. No late assignments will be accepted. Strictly individual (no team) solutions are required. 

Place your solution as file `CA1.ipynb` into the folder `AM1054\CA1` of your `Z:` drive. 

Write `std_nr='11XXXXXXX'` (where 11XXXXXXX is your student number) into the next cell:

<!-- grade_id: stnr, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7c82e818061b6421, points: 0, grade: true, solution: false -->
```python
assert(int(std_nr))
```

Write your solutions in the cell below each question. 

# 1. The Henon Sequence

Consider the Henon sequence 

\begin{equation}
((x_0,y_0), (x_1,y_1), (x_2,y_2),\ldots),
\end{equation}
which consists of pairs of numbers, and is iteratively defined as follows:

\begin{align}
x_0 &= 0\\
y_0 &= 0\\
x_{{n+1}}&=1-ax_{n}^{2}+y_{n}
\\y_{{n+1}}&=bx_{n}.
\end{align}
$a$ and $b$ are positive parameters.

a) Write a Python function `henon_sequence(N, a=1.4, b=0.3)` which returns a list of the first $N$ pairs of the Henon sequence for the classical parameters $a=1.4$ and $b=0.3$. Calculate `henon_sequence(10)`. 

<!-- grade_id: 4a, points: 0, grade: false, solution: true -->
```python
def henon_sequence(N, a=1.4, b=0.3):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-cb04af5990e7b859, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(henon_sequence(10)[-1],(1.1831385762027349, -0.08636135161151118))
```

b) Plot the first 1000 pairs of the Henon sequence as red dots in an $x$-$y$ diagram. 

<!-- grade_id: cell-3b7f24ec0699c517, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

c) Redo the last plot, and highlight the pair with the largest distance from the origin with a green square. 

<!-- grade_id: cell-e02b94ee655b0586, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

d) Redo the last plot and add two sliders, which allow you to change `a` in the interval from 0 to 1.4 in steps of 0.02 and `b` from 0 to 0.3 in steps of 0.01.

<!-- grade_id: cell-6fd479b5bb2383ca, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 2 Working with data files

The file "dat.csv" contains three columns of data.  Do the following with this data:

a) Generate a new file "dat_new.csv" with two columns: The first column is identical with the first column of "dat.csv" and the second column is the product of the last two columns of "dat.csv". 

<!-- grade_id: 1ci, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-17ab538b5044910d, points: 1, grade: true, solution: false -->
```python
assert(np.genfromtxt("dat_new.csv", delimiter=",").shape == (101,2))

```

b) Use Python to count the number of rows in `dat.csv`, where the abolute value of the last column is larger than the absolute value in the second column. Store this number in a variable called `count`

<!-- grade_id: 1cii, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f32f14878426735, points: 1, grade: true, solution: false -->
```python
assert(count>0)
```

c) Use one line of code to assign the first 5 rows of `dat.csv` in reverse order to an array named `fiverows`.

<!-- grade_id: 1civ, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0f0de77adb930cf4, points: 0.5, grade: true, solution: false -->
```python
assert(len(fiverows)==5)
```

# Q3. Presentation Skills

In the following links replace `xxx` with the last three digits of your student number.  

1. http://arxiv.org/abs/1810.04XXX
1. http://arxiv.org/abs/1810.03XXX
1. http://arxiv.org/abs/1810.02XXX
1. http://arxiv.org/abs/1810.01XXX
1. http://arxiv.org/abs/1809.09XXX
1. http://arxiv.org/abs/1809.08XXX
1. http://arxiv.org/abs/1809.07XXX
1. ....


a) Open the paper in the first link and look at the section headings of this paper. Reproduce the outline structure of this paper, by copying the section headings into the cell below with the appropriate heading level. Also include title, abstract, references, etc. as appropriate (Note: only copy section headings, not any part of the main text.)

<!-- grade_id: 3a, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

b) Check through the papers in the links above and find mathematical formulas.  Reprodce them in the markdown cell below. Do this for a number of formulas, until you have at least 300 characters of latex code. Precisely state, which formulas you reproduce from which paper. 

<!-- grade_id: 3b, points: 1.5, grade: true, solution: true -->
YOUR ANSWER HERE

c) Check through the papers in the links above and see, if any of them contains any two dimensional plots. If this is the case, choose one such plot, which contains either at least two different lines or at least two subplots. Reprodce the axes of the graph, including axis labels, correct range of axes. Plot some fake data with about 10 data points inside the graphs. The aim is to produce a graph, which at first sight looks roughly similar to the published plot. Precisely state, which graph you reproduce from which paper. 

<!-- grade_id: 3d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```